Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Usama Puri"
COLLABORATORS = "None"
```

---

# CS110 Pre-class Work 2.2

## Question 1 (Exercise 3.1-3 of Cormen, et al. )
Explain why the statement, "The running time of algorithm A is at least $O(n^2)$," is meaningless.


<!-- grade_id: cell-437d61c1420d4f99, points: 0, grade: true, solution: true -->
The O(n^2) bound on the worst-case running time of an algorithm, however, does not imply a ‚O(n^2) bound on the running time of that algorithm for every input. 

For example, when the input is already sorted, insertion sort runs in ‚ O(n) time. Technically, it is an abuse to say that the running time of insertion sort is O(n^2), since for a given n, the actual running time varies, depending on the particular input of size n. When we say “the running time is O(n^2),” we mean that there is a function f(n) that is O(n^2) such that for any value of n, no matter what particular input of size n is chosen, the running time on that input is bounded from above by the value f(n). Equivalently, we mean that the worst-case running time is O(n^2)

## Question 2 (Exercise 3.1-4 of Cormen, et al. )

Is $2^{n+1}=O(2^n)$? Is $2^{2n}=O(2^n)$?

<!-- grade_id: cell-6a0c4ee19dadfddf, points: 0, grade: true, solution: true -->
2^(n+1) ≥ 2(2n) for all n ≥ 0, so 2^(n+1) = O(2^n).

However, 2^2n is not O(2^n). If it were, there would exist n0 and c such that n ≥ n0 implies 2n x 2n = 2^2n ≤ c2n,
so 2n ≤ c for n ≥ n0 which is clearly impossible since c is a constant.

## Question 3.
Write a function in Python that solves the maximum-subarray problem using a brute-force approach. Your Python function must:
* Take as Input an array/list  of numbers
* Produce the following Output: 
    * the start and end indices of the subarray containing the maximum sum.
    * value of the maximum subarray (float)


<!-- grade_id: cell-527e6532d6992fd0, points: 0, grade: false, solution: true -->
```python
def bruteforce_max_subarray(A):
    length=len(A)
    mid=A[int(length/2)]
    low=A[0]
    high=A[-1]
    left_sum = -1000
    summ = 0
    for i in range(mid, low-1, -1):
        summ += A[i]
        if summ > left_sum:
            left_sum = summ
            max_left = i
    right_sum = -1000
    summ = 0
    for j in range(mid+1, high):
        summ += A[j]
        if summ > right_sum:
            right_sum = summ
            max_right = j
    return(max_left, max_right, left_sum + right_sum)
L=[-1, 2, 3, 4, 5, 6, 7, 8, 9, -10]
bruteforce_max_subarray(L)
```

Output:
```
(1, 8, 44)
```

<!-- grade_id: cell-a28a56466c9537ff, points: 1, grade: true, solution: false -->
```python
assert(bruteforce_max_subarray([-2,1,-1,2,-5]) == (1, 3, 2))
assert(bruteforce_max_subarray([-2, -5, 6, -2, -3, 1, 5, -6]) == (2, 6, 7))
```

Output:
```
TypeError                                 Traceback (most recent call last)
<ipython-input-5-5ddab6f2f574> in <module>()
----> 1 assert(bruteforce_max_subarray([-2,1,-1,2,-5]) == (1, 3, 2))
      2 assert(bruteforce_max_subarray([-2, -5, 6, -2, -3, 1, 5, -6]) == (2, 6, 7))

TypeError: bruteforce_max_subarray() missing 3 required positional arguments: 'low', 'mid', and 'high'
```

## Question 4. 
Test your Python maximum-subarray function using the following input list (from Figure 4.3 of Cormen et al.):  
`A = [13, -3, -25, 20, -3, -16, -23, 18, 20, -7, 12, -5, -22, 15, -4, 7] `

If your Python implementation is correct, your code must return: 
* 43 - which is the answer to the maximum subarray problem, and 
* <7, 10> -the start and the end indices of the max subarray. 



<!-- grade_id: cell-e4a632ce0edc1697, points: 0, grade: true, solution: true -->
```python
def bruteforce_max_subarray(A, low, mid, high):
    left_sum = -1000
    summ = 0
    for i in range(mid, low-1, -1):
        summ += A[i]
        if summ > left_sum:
            left_sum = summ
            max_left = i
    right_sum = -1000
    summ = 0
    for j in range(mid+1, high):
        summ += A[j]
        if summ > right_sum:
            right_sum = summ
            max_right = j
    return(max_left, max_right, left_sum + right_sum)
A = [13, -3, -25, 20, -3, -16, -23, 18, 20, -7, 12, -5, -22, 15, -4, 7]
bruteforce_max_subarray(A, 0,8 ,-1)
```

Output:
```
<ipython-input-60-2464841c993d> in <module>()
     16     return(max_left, max_right, left_sum + right_sum)
     17 A = [13, -3, -25, 20, -3, -16, -23, 18, 20, -7, 12, -5, -22, 15, -4, 7]
---> 18 bruteforce_max_subarray(A, 0,8 ,-1)

<ipython-input-60-2464841c993d> in bruteforce_max_subarray(A, low, mid, high)
     14             right_sum = summ
     15             max_right = j
---> 16     return(max_left, max_right, left_sum + right_sum)
     17 A = [13, -3, -25, 20, -3, -16, -23, 18, 20, -7, 12, -5, -22, 15, -4, 7]
     18 bruteforce_max_subarray(A, 0,8 ,-1)

UnboundLocalError: local variable 'max_right' referenced before assignment
```

## Question 5. Asymptotic notation. 
Complete the following table using the asymptotic notation that best describes the problem. For example, if both $O(n^3)$ and $O(n)$ are possible for an algorithm, the answer is $O(n)$ because the function $f(n) = O(n)$ provides a tighter and more accurate fit; if both $O(n)$ and $\Theta(n)$ are possible, the correct answer is $\Theta(n)$ because $\Theta(n)$ provides both information about the upper and lower bound, thus it is more accurate than $O(n)$.

You should copy the following table and paste and edit it in the cell below. 

Algorithm | Big Oh ($O$) | Big Theta ($\Theta$)
--- | --- | ---
Insertion sort |  |
Selection sort |  |
Bubble sort |  | 
Finding maximum subarray |  |

<!-- grade_id: cell-247c51df276b622e, points: 0, grade: true, solution: true -->

Algorithm | Big Oh ($O$) | Big Theta ($\Theta$)
--- | --- | ---
Insertion sort | O(n^2) | O(n^2))
Selection sort | O(n^2) |  O(n^2)
Bubble sort | О(n^2) | O(n)
Finding maximum subarray |O(n^2)  |

## [Optional] Question 6. 
How can you change this code to make it find the minimum-subarray?

<!-- grade_id: cell-f4553f71858d1bc4, points: 0, grade: false, solution: true -->
```python
def bruteforce_min_subarray(A):
    """Implements brute-force maximum subarray finding.
    
    Inputs:
    - A: a NON-EMPTY list of floats
    
    Outputs: A tuple of
    - the start index of the max subarray
    - the end index of the max subarray
    - the value of the maximum subarray
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-4263c7494a5f09d3, points: 0, grade: true, solution: false -->
```python
assert(bruteforce_min_subarray([1]*10)[0] == bruteforce_min_subarray([1]*10)[1])
assert(bruteforce_min_subarray([1]*10)[2] == 1)
```

```python


```