Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# CS110 Pre-class Work 4.2

## Part A. The Hire-Assistant Problem.

Imagine that you need to hire a new assistant. Every day an agency sends a new assistant for you to interview. If the assistant is better than your current assistant, then you fire your current assistant and you hire the better assistant. You may assume that assistant quality is uniformly distributed between 0 and 1.

## Question 1.
Write a function, named hire_assistant, that takes applicants (a list of the numbers that represent the level of qualification of the applicants; the higher the number, the better qualified), and returns the number hires if the applicants are presented in the exact same order as the input list applicants. Note that your function should not randomize anything (or else it would be called a randomized algorithm).

<!-- grade_id: cell-d011f5f4707fe41a, points: 0, grade: false, solution: true -->
```python
def hire_assistant(applicants):
    hired = []
    best = 0
    for i in applicants:
        if i>best:
            best = i
            hired.append(i)
    return hired
hire_assistant([1, 24, 3, 43, 51, 6, 3, 92, 11, 87])
```

Output:
```
[1, 24, 43, 51, 92]
```

<!-- grade_id: cell-66778b97ad66f71e, points: 1, grade: true, solution: false -->
```python
assert(hire_assistant([1])==1)
assert(hire_assistant([-1, -2, -3, -4])==1)
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-10-2dddc1aba31f> in <module>()
----> 1 assert(hire_assistant([1])==1)
      2 assert(hire_assistant([-1, -2, -3, -4])==1)

AssertionError:
```

## Question 2. 
Assuming the applicants are presented in a random order, write a function that receives the number of applicants as input and returns the average number of assistants hired.

**N.B.:** Don’t forget to run the simulation several times for each given number of applicants to better estimate the number of hires (please refer to task 3 of the Study Guide).

<!-- grade_id: cell-b223520ca72942a0, points: 0, grade: true, solution: true -->
```python
import random
def experimental_hires(n):
    hired = []
    lixt=[]
    avg=[]
    best = 0
    for i in range (1, n):
        r=random.randint(1,n)
        lixt.append(r)
    for i in lixt:
        if i>best:
            best=i
            hired.append(best)
        avg.append(sum(hired))
    sume=sum(hired)
    length=len(hired)
    average=sume/length
    print (average)
    return hired
experimental_hires(100)

```

Output:
```
95.0
[91, 99]
```

```python
def averg():
    for i in range(1, 10):
        experimental_hires(100)
    lenth=len(avg)
    summe=sum(avg)
    average=summe/lenth
    print(average)

averg()
```

Output:
```
... last 1 frames repeated, from the frame below ...

<ipython-input-13-0c611d735dec> in experimental_hires(n)
     14         avg.append(len(hired))
     15     for i in range(1, 10):
---> 16         experimental_hires(100)
     17     lenth=len(avg)
     18     summe=sum(avg)

RecursionError: maximum recursion depth exceeded while calling a Python object
```

## Question 3.

Use the function below, `analytical_hires(N)`, which returns the analytical expected number of hires, given the number of applicants, along with the function you created in question 2 to create a graph with two curves such that:
* The x-axis shows the total number of applicants (make sure label the x-axis)
* The y-axis shows the average number of hires (make sure label the y-axis)
* The graph contains two curves;
    * Curve 1: the theoretical performance estimates computed calls to the function `analytical_hires`.
    * Curve 2: the simulated or experimental estimates using the function you created in question 2.


```python
def analytical_hires(N):
    """
    Return the analytical expected number of hires if there are N applicants
    Inputs:
    - N: Number of applicants
    Outputs:
    - hires: Average number of assistants hired
    """
    # from the textbook, we know that the analytical result is 
    # 1 + 1/2 + 1/3 + ... + 1/N
    hires = 0
    for n in range(N):
        hires += 1/(n+1)
    return hires
```

<!-- grade_id: cell-f9c07920c069ce20, points: 0, grade: true, solution: true -->
```python
analytical_hires(100)

# plot the hiring candidates agaisnt the numbero all applicants
Z1, Y1 = [], []
for i in range(10, 10010, 1):
    Y1.append(len(analytical_hires(i)))

Z1 = np.mean(np.array(Y1).reshape(-1, 50), axis=1)    

plt.figure(figsize=(10, 8))
plt.plot(Z1) #, alpha = 0.8, s=7, marker='o')
plt.xlabel("Number of candidates")
plt.title("Candidates hired")
# plt.ylim(0, 30)
plt.show()
```

## Question 4.

Plot a graph with the x-axis showing the total number of applicants and the y-axis showing the probability that exactly one assistant is hired.

<!-- grade_id: cell-d3fe1b7d6d175ad7, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## [Optional] Question 5.
Assume that an assistant is able to perform an amount of work each day that is equal to their “quality”. You have a total amount of work M that needs to be accomplished. Your costs are as follows:
* X = daily salary for the assistant,
* Y = fee to the employment agency,
* Z = retrenchment fee for the old assistant.

Try to formulate an optimal stopping rule (ie. at what point should one stop requesting new potential hires from the agency?) Make any necessary assumptions to ensure the problem is well-formulated.


<!-- grade_id: cell-af2f0291eced6982, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Part B. The Hat Check Problem.

There is a coat check at a party, where an attendant stores everyone’s hat while they attend the party. The attendant receives the N hats from everyone attending (all attendees come with a hat). Unfortunately, the coat check attendant forgets which hat belongs to whom. Rather than admitting a mistake, the attendant simply returns random hats back to the party goers. 
What is the average number of correct hats returned? Here are some guiding questions to help you to simulate this problem. 

## Question 1. 
Knowing that everyone’s hats are unique and every guest has a hat. Do you need to generate a random sample in a similar way as what you did for the hiring assistant problem? 

<!-- grade_id: cell-e786799fc4eb1499, points: 0, grade: true, solution: true -->
Yes, In my view we can generate a random sample to generate the list of guesses on number of people who will receive their hats. Then We can just calculate the probability of these hats given that we have the sample size of all the owners

## Question 2. 
Which of the following commands do you think is the Pythonic way to implement that? 
```
import numpy as np
n = 100 #the number of party attendants `
```
**Command 1. **
```
hat_list = [np.random.integers(0,n) for i in range(n)]`
```
**Command 2.**
```
hat_list = list(range(n)) 
np.random.shuffle(hat_list) 
```
**Command 3.**
```
hat_list = np.random.sample(n)
```

<!-- grade_id: cell-b8da78e72c1c0738, points: 0, grade: true, solution: true -->
I will prefer using Commang 3 as it generates a list of random integers from the range provided to it. After obtaining the list we can calculate the probability using it as a sample size 

## Question 3.
Now write a function `hat_check(N)` that has: 
* Input: N the number of party attendants. 
* Output: the number of hats correctly returned despite the fact that hats are randomly handed back to the guests.

You should use the command you picked for question 2. 

<!-- grade_id: cell-c8499aeb1b1d76c7, points: 0, grade: true, solution: true -->
```python
def hats(n):
    success = 0
    guests = list(range(1, n+1))
    mis_hat = random.sample(guests, len(guests))
    for i in range(1, n):
        if guests[i] == mis_hat[i]:
            success+=1
    return success/n


```

## Question 4.

Plot a curve with the x-axis showing the total number of party attendants and the y-axis showing the average number of hats correctly returned. As always, remember to run several trials. 

<!-- grade_id: cell-597031ea2a5a512a, points: 0, grade: true, solution: true -->
```python
import random
import matplotlib.pyplot as plt
X, Y = [], []
for i in range(5, 500, 10):
    X.append(i)
    Y.append(hats(i))

plt.figure(figsize=(10, 8))
plt.plot(X, Y) #, alpha = 0.8, s=7, marker='o')
plt.xlabel("Number of guests")
plt.title("Accuracy of hat return")
plt.show()
```

Output:
```
<Figure size 1000x800 with 1 Axes>
```

## [Optional] Question 5.
As $N$ tends to infinity, the number of correct hats returned tends towards a well-known statistical distribution. State the distribution with all its parameters. Plot several samples using your code. Does the empirical distribution match your theoretical prediction?

<!-- grade_id: cell-32fe26c1d99fdd2a, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE