Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# CS110 Pre-class Work 5.2

## Question 1.
Using Figure 7.1 in Cormen et al. as a model, perform manually the partition process on the following list: A = [1,5,6,2,3,8,9,4,7]. You just need to specify the followings:
1. The array after the process is done.
2. The value of $i$ after the process is done.

<!-- grade_id: cell-7aa520f8af13679b, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 2.
Code up a Python implementation of `partition(A, p, r)`, closely follow the pseudo-code in Cormen et al., p.172. Your function should return the index of the pivot in the array.

<!-- grade_id: cell-1ceb2600756c60ff, points: 0, grade: false, solution: true -->
```python
import random
import time
random.seed(123)

# creating a quicksort algorithm
def partition(A, p, r):
    x = A[r]
    i = p-1
    for j in range(p, r):
        if A[j] <= x:
            i+=1
            A[i], A[j] = A[j], A[i]
    A[i+1], A[r] = A[r], A[i+1]
    return i+1
```

<!-- grade_id: cell-a57b60117a7b82fb, points: 1, grade: true, solution: false -->
```python
A = [1,5,6,2,3,8,9,4,7]
assert(partition(A, 0, len(A)-1)==6)
```

## Question 3.

Code up your own Python implementation of `quicksort(A, p, r)`, using `partition(A,p,r)`.

<!-- grade_id: cell-8c39ebb8cd1aa83a, points: 0, grade: false, solution: true -->
```python
def quick_sort(A, p, r):
    if  p < r:
        q = partition(A, p, r)
        quick_sort(A, p, q-1)
        quick_sort(A, q+1, r)
```

<!-- grade_id: cell-4f822430efd456ee, points: 1, grade: true, solution: false -->
```python
A = [0]
assert(quick_sort(A, 0, 0) == [0])
A = [3,1,2]
assert(quick_sort(A, 0, 2) == [1,2,3])
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-16-3ff2a53494de> in <module>()
      1 A = [0]
----> 2 assert(quick_sort(A, 0, 0) == [0])
      3 A = [3,1,2]
      4 assert(quick_sort(A, 0, 2) == [1,2,3])

AssertionError:
```

## Question 4. 
Explain (using experimental plots) the running time of `quick_sort` when: 
1. all elements of array A have the same value (e.g., [1,1,1])?
2. array A contains distinct elements sorted in decreasing order (e.g., [5,4,2,1])?


<!-- grade_id: cell-b58035dd5fa02329, points: 0, grade: true, solution: true -->
```python
import time
import matplotlib.pyplot as plt

lst1 = []
lst2 = []
x=1000
# trying the algorithm on list 1
for i in range (1,1000):
    y=1
    lst1.append(y)
start1 = time.time()
quick_sort(lst1, 0, len(lst1)-1)
end1 = time.time()
t_time=end1-start1
print('Runnig time random list: ', end1-start1)

# trying the algorithm on list 2

while x!=0:
    lst2.append(x)
    x=x-1

start2 = time.time()
quick_sort(lst2, 0, len(lst2)-1)
end2 = time.time()
t2_time=end2-start2
print('Runnig time ordred list: ', end2-start2)

for x in range(1,1000):
    plt.plot(x,t_time)
    plt.plot(x,t2_time)
    plt.title('Runtime of Qsort')
    plt.xlabel('Total Numbers Sorted')
    plt.ylabel('Seconds')
    plt.legend(['With same integer', 'With Decreasing integers'], loc = 'upper left')
plt.show()
```

```python

```