Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Usama Puri"
COLLABORATORS = "Youtube, peer guidance"
```

---

# CS110 Pre-class Work 3.1

## Question 1. 
Paste in your Python implementation of the maximum subarray from the previous class in the cell below and use that to find out the value of the maximum subarray of this array: `A = [-2, -3, 4, -1, -2, 1]`

<!-- grade_id: cell-e69c8bbcd40ca242, points: 0, grade: true, solution: true -->
```python
def max_subarray(A, low, mid, high):
    left_sum = -1000
    summ = 0
    for i in range(mid, low-1, -1):
        summ += A[i]
        if summ > left_sum:
            left_sum = summ
            max_left = i
    right_sum = -1000
    summ = 0
    for j in range(mid+1, high):
        summ += A[j]
        if summ > right_sum:
            right_sum = summ
            max_right = j
    return(max_left, max_right, left_sum + right_sum)
L= [-2, -3, 4, -1, -2, 1]

max_subarray(L, 0, 3, 6)
```

Output:
```
(2, 5, 2)
```

## Question 2. 
Now, your friend Joe comes and appends a single extra number at the end of the array, which becomes:  `B = [-2, -3, 4, -1, -2, 1, 8]`. Do you have to re-run the entire maximum subarray again? Explain your answer. 
The subsequent questions will help you figure out an efficient algorithmic strategy to address the last question, but make sure to write your explanation above first, before answering the remaining questions.


```python
def max_subarray(A, low, mid, high):
    left_sum = -1000
    summ = 0
    for i in range(mid, low-1, -1):
        summ += A[i]
        if summ > left_sum:
            left_sum = summ
            max_left = i
    right_sum = -1000
    summ = 0
    for j in range(mid+1, high):
        summ += A[j]
        if summ > right_sum:
            right_sum = summ
            max_right = j
    return(max_left, max_right, left_sum + right_sum)
L= [-2, -3, 4, -1, -2, 1,8]
max_subarray(L, 0, 3, 7)
```

Output:
```
(2, 6, 10)
```

<!-- grade_id: cell-d4b7cd845c816263, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 3. 

**Determine if the following statement is True or False and explain your answer.**
If the maximum subarray of the array A is different than the maximum subarray of the array B (questions 1 and 2), the new maximum-subarray doesn’t need to include 8 (i.e., the newly appended element). 


<!-- grade_id: cell-9377964a8756f13b, points: 0, grade: true, solution: true -->
This statement is true. As the maximum subarray is re calculated by the bruteforce Algorithm and all possible sums are checked so if the incrementation of another number lowers the max sum it wont be added into the max sum of the array


## Question 4.
Complete the Python function `incremental_max_subarray(x, mx)` in the cell below.

This [video](https://www.youtube.com/watch?v=AAgErqQmwmA&list=PLF_a-qBXTGFektoI6JUOTRL36JlvD04BR&index=4&t=0s)  might be helpful to understand the `incremental_max_subarray` problem.

<!-- grade_id: cell-0230e459f3d701f8, points: 0, grade: false, solution: true -->
```python
def max_subarray_sum(L):
    max_so_far=L[0]
    curr_max = L[0]
    for i in range (1, len(L)):
        curr_max = max(L[i], curr_max + L[i])
        max_so_far = max(max_so_far, curr_max)
    return max_so_far

def incremental_max_subarray(x, mx):
    curr_max = max(mx, mx + x[-1])
    return curr_max


# Testing the algorithm
B = [-2, -3, 4, -1, -2, 1]
max_subarray_sum(B)
# OUTPUT = 2


# Appending an extra positive value at the end of the list
B.append(8)
incremental_max_subarray(B, 2)
# OUTPUT = 10
```

Output:
```
10
```

<!-- grade_id: cell-9abf19e56620ffa3, points: 1, grade: true, solution: false -->
```python
B = [-2, -3, 4, -1, -2, 1, 8]
assert(incremental_max_subarray(B, 4) == 10)
assert(incremental_max_subarray(B[:1], 0) == B[0])
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-55-2256e7bbc067> in <module>()
      1 B = [-2, -3, 4, -1, -2, 1, 8]
----> 2 assert(incremental_max_subarray(B, 4) == 10)
      3 assert(incremental_max_subarray(B[:1], 0) == B[0])

AssertionError:
```

<!-- grade_id: cell-bacebe71191cbd0f, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: cell-fd96d4ccccd99fe6, points: 0, grade: false, solution: true -->
```python
"""
    Using `incremental_max_subarray` iteratively on A to produce the value of the maximum
    subarray of A.
    
    Inputs:
    - A: a NON-EMPTY list of floats
    
    Outputs: flaot, the sum of the maximum subarray of A
"""
def max_subarray_sum(L):
    max_so_far=L[0]
    curr_max = L[0]
    for i in range (1, len(L)):
        curr_max = max(L[i], curr_max + L[i])
        max_so_far = max(max_so_far, curr_max)
    return max_so_far

def incremental_max_subarray(x, mx):
    curr_max = max(mx, mx + x[-1])
    return curr_max


# Testing the algorithm
B = [-2.5, -3.5, 4.25, -1.5, -2.25, 1.5]
max_subarray_sum(B)
# OUTPUT = 2.25


# Appending an extra positive value at the end of the list
B.append(8.75)
incremental_max_subarray(B, 2.25)
# OUTPUT = 11.0

```

Output:
```
11.0
```

<!-- grade_id: cell-669ad779766aa2c3, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

### Part 2.
Is this more efficient than the divide-and-conquer approach? Explain.

<!-- grade_id: cell-cd8f0b130a7136db, points: 0, grade: true, solution: true -->
To figure out which Algorithm is the more efficient one we have to consider the number of steps taking place in the Algorithms and we can then guess their Big O notation using recursion tree method.

```python

```